# STATS 507, Week 1 practice

```python
type(3)

type(3.0)

type(True)

type(None)

"I like cats. " * 3 
"I like cats. " == 3 
```

Output:
```
False
```

### Problem 1
Write a Python function that prints "Hello, World!".

```python
def print_hello_world(name):
    print("Hello World, " + name + "!")

print_hello_world("Zhu")
```

Output:
```
Hello World, Zhu!
```

```python
print("Hello World")
```

Output:
```
Hello World
```

### Problem 2
Write a function that takes two integers as argumens. Return the first argument if it is a positive number. Otherwise, return the sum of both arguments.

<!-- grade_id: cell-db7233a2d95c6406, points: 0, grade: false, solution: true -->
```python
def postive_or_sum(x, y):
    if x >= 0: 
        return x
    return x + y

assert postive_or_sum(1,2) == 1
```

<!-- grade_id: cell-d525ecbed5402db1, points: 1, grade: true, solution: false -->
```python
assert postive_or_sum(42, 16) == 42
assert postive_or_sum(33, 13) == 33
assert postive_or_sum(-33, 13) == -20
```

```python
assert postive_or_sum(3,4) == 7, "Wrong result for 3,4"
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[17], line 1
----> 1 assert postive_or_sum(3,4) == 7

AssertionError:
```